<div class="alert alert-danger" style="color:black"><b>Running ML-LV Jupyter Notebooks:</b><br>
    <ol>
        <li>Make sure you are running all notebooks using the <code>adv_ai</code> kernel.
        <li><b>It is very important that you do not create any additional files within the weekly folders on CSCT cloud.</b> Any additional files, or editing the notebooks with a different environment may prevent submission/marking of your work.</li>
            <ul>
                <li>NBGrader will automatically fetch and create the correct folders files for you.</li>
                <li>All files that are not the Jupyter notebooks should be stored in the 'ML-LV/data' directory.</li>
            </ul>
        <li>Please <b>do not pip install</b> any python packages (or anything else). You should not need to install anything to complete these notebooks other than the packages provided in the Jupyter CSCT Cloud environment.</li>
    </ol>
    <b>If you would like to run this notebook locally you should:</b><br>
    <ol>
        <li>Create an environment using the requirements.txt file provided. <b>Any additional packages you install will not be accessible when uploaded to the server and may prevent marking.</b></li>
        <li>Download a copy  of the notebook to your own machine. You can then edit the cells as you wish and then go back and copy the code into/edit the ones on the CSCT cloud in-place.</li>
        <li><b>It is very important that you do not re-upload any notebooks that you have edited locally.</b> This is because NBGrader uses cell metadata to track marked tasks. <b>If you change this format it may prevent marking.</b></li>
    </ol>
</div>

# 2 Introduction to NBGrader

Throughout this module we will be using Jupyter Notebooks containing:
- Further worked examples and activities to illustrate and provide ‘hands-on’ experience of the algorithms and issues covered in that week’s materials.
- A series of exercises to complete to demonstrate your understanding. These will involve a mixture of:
    - *Automatically* marked exercises which involve writing or completing code to implement ML algorithms or functions to implement create common exercises from typical Data Science/AI workflows.
    - *Manually* marked exercises/questions which will contain a mixture of short essay style questions, multiple choice or short answer questions and some larger coding related exercises.

We will be using NBGrader to *fetch*, *validate* and *submit* each notebook for marking, and receive *feedback*. For a complete overview of this process you should refer to the assessment brief and instructions provided by your tutor. The following exercises provide a very simple demonstration for how the automatic and manually marked exercises work within NBGrader.

<div class="alert alert-success" style="color:black"><b>Note:</b> The following exercises <b>do not</b> contribute towards your portfolio. They are only intended to demonstrate the process of completing exercises using NBGrader. As such they are worth <b>0 Marks</b>.
</div>

<div class="alert alert-info" style="color:black"><h3>2.1 Exercise: Automatically marked</h3>

With NBGrader automatically marked cells simply check that your code passes a series of tests. These could involve checking variable values, or completing a function and checking it produces the expected output or error handling. There are two types of test cells:
- Visible - Most of the tests will be visible to you, like the 5 in the cell below.
- Hidden - Some of the tests will not be visible to you. Your code will only be run against these tests once it has been submitted for marking via NBGrader.

All exercises (highlighted in blue cells like this one) will include a description of what your code should do and the total marks available. Ensure you complete any code marked with `# YOUR CODE HERE`.

You can check your work by clicking on the 'Validate' button at the top of the page (in JupyterHub), or simply ensure all cells run without producing an error.

For this exercise you are asked to write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a **ValueError**. For example, your function should print **[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]** for $n=10$.

<b>MARKS AVAILABLE: 0</b>
</div>

<!-- grade_id: cell-728d2fe8b7067637, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    """
    # YOUR CODE HERE
    if n <= 0:
        raise ValueError("n must be a positive integer.")
    return [i**2 for i in range(1, n+1)]
```

<!-- grade_id: cell-1a8de0a6dee68aff, points: 0, grade: true, solution: false -->
```python
# Check that squares returns the correct output for several inputs
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(3) == [1, 4, 9]
assert squares(4) == [1, 4, 9, 16]
assert squares(5) == [1, 4, 9, 16, 25]

# Check that squares raises an error for invalid inputs
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

<div class="alert alert-info" style="color:black"><h3>2.2 Exercise: Manually marked</h3>

With NBGrader manually marked cells must be marked by your tutor once the notebook has been submitted for marking via NBGrader. These can be either markdown cells or code cells. Markdown cells will be used for short essay style questions, multiple choice or short answer questions and code cells for some larger coding related exercises.

All manually marked exercises will include a description of what your code should do and the marks available, including a marking scheme. Ensure you complete any cell marked with `"YOUR ANSWER HERE"`.

The following is an example of one short answer and one multiple choice question.

</div>

<div class="alert alert-info" style="color:black"><h3>2.2.1 Short Answer:</h3>

Simply state the name of one of the films in your reviews dataset and whether you think the overall sentiment (from your annotations) is positive or negative. For example, "Batman - Positive".

You can complete the question by editing the markdown cell below.

<b>MARKS AVAILABLE: 0</b>
</div>

<!-- grade_id: cell-705906757aed343a, points: 0, grade: true, solution: true -->
Dr. Strange - Positive
XO, Kity - Negative

<div class="alert alert-info" style="color:black"><h3>2.2.2 Multiple Choice:</h3>

How much was Google fined under GDPR laws for dropping cookies on google.fr (hint: you may want to read the article about web scraping linked above)?

<ol type="A">
    <li>€120,000</li>
    <li>€50,000,000</li>
    <li>€120,000,000</li>
</ol>

You can complete the question by editing the markdown cell below.

<b>MARKS AVAILABLE: 0</b>
</div>

<!-- grade_id: cell-89d11d4b233fd832, points: 0, grade: true, solution: true -->
€50,000,000

<div class="alert alert-success" style="color:black"><h3>Before you submit this notebook to NBGrader for marking:</h3> 

1. Make sure have completed all exercises marked by <span style="color:blue">**blue cells**</span>.
2. For automatically marked exercises ensure you have completed any cells with `# YOUR CODE HERE`. Then click 'Validate' button above, or ensure all cells run without producing an error.
3. For manually marked exercises ensure you have completed any cells with `"YOUR ANSWER HERE"`.
4. Ensure all cells are run with their output visible.
5. Fill in your student ID (**only**) below.
6. You should now **save and download** your work.

</div>

**Student ID:** 15006280